```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

# Lab 10: SQL

In this lab, we are going to practice viewing, sorting, grouping, and merging tables with SQL. We will explore two datasets:
1. A "minified" version of the [Internet Movie Database](https://www.imdb.com/interfaces/) (IMDb). This SQLite database (~10MB) is a tiny sample of the much larger database (more than a few GBs). As a result, disclaimer that we may get wildly different results than if we use the whole database!

1. The money donated during the 2016 election using the [Federal Election Commission (FEC)'s public records](https://www.fec.gov/data/). You will be connecting to a SQLite database containing the data. The data we will be working with in this lab is quite small (~16MB); however, it is a sample taken from a much larger database (more than a few GBs).


The on-time deadline is **Tuesday, November 7th at 11:59 PM PT**. Please read the syllabus for the grace period policy. No late submissions beyond the grace period will be accepted.

### Lab Walk-Through
In addition to the lab notebook, we have also released a prerecorded walk-through video of the lab. We encourage you to reference this video as you work through the lab. Run the cell below to display the video.

**Note:** the walkthrough video is from Spring 2022, where the format of answers was different than it is this semester.

```python
from IPython.display import YouTubeVideo
YouTubeVideo("uQ3E4pejmD8", list = 'PLQCcNQgUcDfpdBnhS-lPq8LPas48tkMgp', listType = 'playlist')
```

### Collaboration Policy
Data science is a collaborative activity. While you may talk with others about this assignment, we ask that you **write your solutions individually**. If you discuss the assignment with others, please **include their names** in the cell below.

**Collaborators:** *list names here*

```python
# Run this cell and the next one to set up your notebook.
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
#import plotly.express as px
import sqlalchemy
from pathlib import Path
from zipfile import ZipFile

with ZipFile('data.zip', 'r') as zipObj:
   # Extract all the contents of zip file in current directory
   zipObj.extractall()
```

```python
# Run this cell to set up SQL. You can ignore any 
# yellow messages that pop up when running this cell
%load_ext sql
```

```python
# Run this cell to set up the autograder
from sql_utils import * 
from pathlib import Path
Path("results").mkdir(parents=True, exist_ok=True)
```

## SQL Query Syntax

Throughout this lab, you will become familiar with the following syntax for the `SELECT` query:

```
SELECT <column list>
FROM <table>
[WHERE <predicate>]
[GROUP BY <column list>]
[HAVING <predicate>]
[ORDER BY <column list>]
[LIMIT <number of rows>]
[OFFSET <number of rows>]
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 0 [Tutorial]: Writing SQL in Jupyter Notebooks

### 1. `%%sql` cell magic

In lecture, we used the `sqlalchemy` extension to use **`%%sql` cell magic**, which enables us to connect to SQL databases and issue SQL commands within Jupyter Notebooks.

Run the below cell to connect to a mini IMDb database.

```python
from pathlib import Path
import shutil
import pathlib

# Do not modify following lines.

# These lines specify the location of our database file
imdbfile = 'imdbmini.db'
tmpimdb = Path('.') / imdbfile

# Specify the database connection path (in this case, a sqlite database in a file)
imdb_sqlite_conn = 'sqlite:///' + str(tmpimdb)

# Create an SQLAlchemy engine to connect to the database
imdb_engine = sqlalchemy.create_engine(imdb_sqlite_conn)
```

The above cell connects to the same database using the SQLAlchemy Python library, which can connect to several different database management systems, including sqlite3, MySQL, PostgreSQL, and Oracle. The library also supports an advanced feature for generating queries called an [object relational mapper](https://docs.sqlalchemy.org/en/latest/orm/tutorial.html) or ORM, which we won't discuss in this course but is quite useful for application development.

```python
# Plug the SQLAlchemy engine into line magic to establish a connection
%sql imdb_engine
```

Above, prefixing our single-line command with `%sql` means that the entire line will be treated as a SQL command (this is called "line magic"). In this class we will most often write multi-line SQL, meaning we need "cell magic", where the first line has `%%sql` (note the double `%` operator).

The database `imdbmini.db` includes several tables, one of which is `Name`. Running the below cell will return first 5 lines of that table. Note that `%%sql` is on its own line.

We've also included syntax for single-line comments, which are surrounded by `--`, and multi-line comments, which are surrounded by `/*` and `*/`.

```python
%%sql
/*
 * This is a
 * multi-line comment.
 */
-- This is a single-line/inline comment. --
SELECT *
FROM Name
LIMIT 5;
```

<br/>

### 2. The `pandas` command `pd.read_sql`

This section describes how data scientists use SQL and `python` in practice, using the `pandas` command `pd.read_sql` ([documentation](https://pandas.pydata.org/docs/reference/api/pandas.read_sql.html)). **You will see both `%sql` magic and `pd.read_sql` in this course**.

With the SQLAlchemy object `engine`, we can call `pd.read_sql` which takes in a `query` **string**. Note the `"""` to define our multi-line string, which allows us to have a query span multiple lines. The resulting `DataFrame` `df` stores the results of the same SQL query from the previous section.

```python
# Run this cell to see the demo.
query = """
SELECT *
FROM Name
LIMIT 5;
"""

df = pd.read_sql(query, imdb_engine)
df
```

#### `pd.read_sql` vs. `%sql` magic error messages

`pd.read_sql` has **long error messages**: Given that the SQL query is now in the string, the errors become more unintelligible. Consider the below (incorrect) query.

**Note**: Uncomment the below code and check out the error. You can uncomment/comment out multiple cells at the same time by selecting the lines and press ctrl+/ or command+/ 

```python
# Uncomment the below code and check out the error.
# query = """
# SELECT *
# FROM Title;
# LIMIT 5
# """
# pd.read_sql(query, imdb_engine)
```

<br/>
<details>
<summary>Now that's an unruly error message! Can you see what's wrong in the cell above and correct the query? Toggle this cell to check your answer!</summary>
It has a semicolon in the wrong place!
</details>

<br/>

On the other hand, `%sql` magic gives more intelligible error messages, so we will use this format more often.

```python
# %%sql
# -- Uncomment the code and check out the error. --
# SELECT *
# FROM Title;
# LIMIT 5
```

<br/>
<!-- 
TODO

### 3. A suggested workflow for writing SQL in Jupyter Notebooks

Which approach is better, `%sql` magic or `pd.read_sql`?

The SQL database generally contains much more data than what you would analyze in detail. As a `python`-fluent data scientist, you will often query SQL databases to perform initial exploratory data analysis, a subset of which you load into `python` for further processing.

In practice, you would likely use a combination of the two approaches. First, you'd try out some SQL queries with `%sql` magic to get an interesting subset of data. Then, you'd copy over the query into a `pd.read_sql` command for visualization, modeling, and export with `pandas`, sklearn, and other `python` libraries.

For SQL assignments in this course, to minimize unruly error messages while maximizing `python` compatibility, we suggest the following "sandboxed" workflow:
1. Create a `%%sql` magic cell using the code below. 

    ```
    %% sql
    -- This is a comment. Put your code here... --
    ```
    
<br/>

2. Work on the SQL query in the `%%sql` cell; e.g., `SELECT ... ;`
3. Then, once you're satisfied with your SQL query, copy it into the multi-string query in the answer cell (the one that contains the `pd.read_sql` call).

You don't have to follow the above workflow to get full credit on assignments, but we suggest it to reduce debugging headaches. We've created the scratchwork `%%sql` cells for you in this assignment—but **if you add any cells, please delete them before submitting**. Newly added cells will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file. -->


<br/><br/>
<hr style="border: 1px solid #fdb515;" />

# Part 1: The IMDb (mini) Dataset

Let's explore a miniature version of the [IMDb Dataset](https://www.imdb.com/interfaces/). This is the same dataset that we will use for the upcoming homework. We'll load it in using cell magic.

```python
imdb_engine = sqlalchemy.create_engine(imdb_sqlite_conn)
```

```python
%sql imdb_engine
```

```python
%%sql
SELECT * FROM sqlite_master WHERE type='table';
```

From running the above cell, we see the database has 4 tables: `Name`, `Role`, `Rating`, and `Title`.

<details open>
    <summary>[<b>Click to Expand</b>] See descriptions of each table's schema.</summary>
    
**`Name`** – Contains the following information for names of people.
    
- nconst (integer) - alphanumeric unique identifier of the name/person
- primaryName (text)– name by which the person is most often credited
- birthYear (text) – in YYYY format
- deathYear (text) – in YYYY format
    
    
**`Role`** – Contains the principal cast/crew for titles.
    
- tconst (integer) - alphanumeric unique identifier of the title
- ordering (text) – a number to uniquely identify rows for a given tconst
- nconst (integer) - alphanumeric unique identifier of the name/person
- category (text) - the category of job that person was in
- characters (text) - the name of the character played if applicable, else '\\N'
    
**`Rating`** – Contains the IMDb rating and votes information for titles.
    
- tconst (integer) - alphanumeric unique identifier of the title
- averageRating (text) – weighted average of all the individual user ratings
- numVotes (text) - number of votes (i.e., ratings) the title has received
    
**`Title`** - Contains the following information for titles.
    
- tconst (integer) - alphanumeric unique identifier of the title
- titleType (text) -  the type/format of the title
- primaryTitle (text) -  the more popular title / the title used by the filmmakers on promotional materials at the point of release
- isAdult (text) - 0: non-adult title; 1: adult title
- Year (text) – represents the release year of a title.
- runtimeMinutes (text)  – primary runtime of the title, in minutes
    
</details>

<br/><br/>

From the above descriptions, we can conclude the following:
* `Name.nconst` and `Title.tconst` are primary keys of the `Name` and `Title` tables, respectively.
* that `Role.nconst` and `Role.tconst` are **foreign keys** that point to `Name.nconst` and `Title.tconst`, respectively.

<br/><br/>

---

## Question 1

What are the different kinds of `titleType`s included in the `Title` table? Write a query to find out all the unique `titleType`s of films using the `DISTINCT` keyword.  (**You may not use `GROUP BY`.**)

```python
# run this cell to make sure we're using the IMDB dataset
%sql imdb_engine 
```

```python
%%sql --save query_1 result_1 <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_1 = %sqlcmd snippets query_1
save_results("result_1", query_1, result_1)
result_1
```

```python
grader.check("q1")
```

<br><br>

---

## Question 2

Before we proceed we want to get a better picture of the kinds of jobs that exist.  To do this examine the `Role` table by computing the number of records with each job `category`.  Present the results in descending order by the total counts.

The top of your table should look like this (however, you should have more rows):

| |category|total|
|-----|-----|-----|
|**0**|actor|21665|
|**1**|writer|13830|
|**2**|...|...|

```python
# run this cell to make sure we're using the IMDB dataset
%sql imdb_engine 
```

```python
%%sql --save query_2 result_2 <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_2 = %sqlcmd snippets query_2
save_results("result_2", query_2, result_2)
result_2
```

```python
grader.check("q2")
```

<br/>
If we computed the results correctly we should see a nice horizontal bar chart of the counts per category below:

```python
# Run this cell to make a bar plot.
res_q2 = result_2.DataFrame() # turn our SQL query to a DataFrame
plt.barh(res_q2["category"], res_q2["total"])
plt.xlabel("Counts")
plt.ylabel("Job Category");
```

<br/><br/>

---

## Question 3

Now that we have a better sense of the basics of our data, we can ask some more interesting questions.

The `Rating` table has the `numVotes` and the `averageRating` for each title. Which 10 films have the most ratings?

Write a SQL query that outputs three fields: the `title`, `numVotes`, and `averageRating` for the 10 films that have the highest number of ratings.  Sort the result in descending order by the number of votes.

**Hint**: The `numVotes` in the `Rating` table is not an integer! Use `CAST(Rating.numVotes AS int) AS numVotes` to convert the attribute to an integer. Sometimes SQL will cast your data automatically when performing operation, but you should cast them explicitly for readability and reproducibility.

```python
# run this cell to make sure we're using the IMDB dataset
%sql imdb_engine 
```

```python
%%sql --save query_3 result_3 <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_3 = %sqlcmd snippets query_3
save_results("result_3", query_3, result_3)
result_3
```

```python
grader.check("q3")
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

# Part 2: Election Donations in New York City

Finally, let's analyze the Federal Election Commission (FEC)'s public records. We connect to the database using cell magic so that we can flexibly explore the database.

```python
# These lines specify the location of our database file
fec_dbfile = 'fec_nyc.db'
fec_tmpdb = Path('.') / fec_dbfile

# Specify SQLAlchemy Connection URL so line/cell magic can talk to database
fec_sqlite_conn = 'sqlite:///' + str(fec_tmpdb)

# Create a SQLAlchemy engine to connect to the database.
fec_engine = sqlalchemy.create_engine(fec_sqlite_conn)

# Plug connection url into line magic to establish connection
%sql fec_engine
```

### Table Descriptions

Run the below cell to explore the **schemas** of all tables saved in the database.

If you'd like, you can consult the below linked FEC pages for the descriptions of the tables themselves.

* `cand` ([link](https://www.fec.gov/campaign-finance-data/candidate-summary-file-description/)): Candidates table. Contains names and party affiliation.
* `comm` ([link](https://www.fec.gov/campaign-finance-data/committee-summary-file-description/)): Committees table. Contains committee names and types.
* `indiv_sample_nyc` ([link](https://www.fec.gov/campaign-finance-data/contributions-individuals-file-description)): All individual contributions from New York City .

```python
%%sql
/* just run this cell */
SELECT sql FROM sqlite_master WHERE type='table';
```

<br/><br/>

Let's look at the `indiv_sample_nyc` table. The below cell displays individual donations made by residents of the state of New York. We use `LIMIT 5` to avoid loading and displaying a huge table.

```python
%%sql
/* just run this cell */
SELECT *
FROM indiv_sample_nyc
LIMIT 5;
```

You can write a SQL query to return the id and name of the first five candidates from the Democratic party, as below:

```python
%%sql
/* just run this cell */
SELECT cand_id, cand_name
FROM cand
WHERE cand_pty_affiliation = 'DEM'
LIMIT 5;
```

<br/><br/>
<hr style="border: 1px solid #fdb515;" />

## [Tutorial] Matching Text with `LIKE`

First, let's look at 2016 election contributions made by Donald Trump, who was a New York (NY) resident during that year. The following SQL query returns the `cmte_id`, `transaction_amt`, and `name` for every contribution made by any donor with "DONALD" and "TRUMP" in their name in the `indiv_sample_nyc` table.

**Notes:**
* We use the `WHERE ... LIKE '...'` to match fields with text patterns. The `%` wildcard represents at least zero characters. Compare this to what you know from regex!
* We use `pd.read_sql` syntax here because we will do some EDA on the result `res`.

```python
# Run this cell to see an example of LIKE.
example_query = """
SELECT 
    cmte_id,
    transaction_amt,
    name
FROM indiv_sample_nyc
WHERE name LIKE '%TRUMP%' AND name LIKE '%DONALD%';
"""

example_res = pd.read_sql(example_query, fec_engine)
example_res
```

If we look at the list above, it appears that some donations were not by Donald Trump himself, but instead by an entity called "DONALD J TRUMP FOR PRESIDENT INC". Fortunately, we see that our query only seems to have picked up one such anomalous name.

```python
# Run this cell to see the value counts for each candidate.
example_res['name'].value_counts()
```

<br/><br/>

---

## Question 4



In the cell below, revise the above query so that the 15 anomalous donations made by "DONALD J TRUMP FOR PRESIDENT INC" do not appear. Your resulting table should have 142 rows. 

**Hints:**
* Consider using the above query as a starting point, or checking out the SQL query skeleton at the top of this lab. 
* The `NOT` keyword may also be useful here.


```python
# run this cell to make sure we're using the FEC dataset
%sql fec_engine 
```

```python
%%sql --save query_4 result_4 <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_4 = %sqlcmd snippets query_4
save_results("result_4", query_4, result_4)
result_4
```

```python
# Print the number of rows in your query
# Double check that this equals 142 
result_4.DataFrame().shape[0]
```

```python
grader.check("q4")
```

<br/><br/>

---

## Question 5: `JOIN`ing Tables

Let's explore the other two tables in our database: `cand` and `comm`.

The `cand` table contains summary financial information about each candidate registered with the FEC or appearing on an official state ballot for House, Senate or President.

```python
%%sql
/* just run this cell */
SELECT *
FROM cand
LIMIT 5;
```

The `comm` table contains summary financial information about each committee registered with the FEC. Committees are organizations that spend money for political action or parties, or spend money for or against political candidates.

```python
%%sql
/* just run this cell */
SELECT *
FROM comm
LIMIT 5;
```

<br><br>

---

### Question 5a

Notice that both the `cand` and `comm` tables have a `cand_id` column. Let's try joining these two tables on this column to print out committee information for candidates.

List the first 5 candidate names (`cand_name`) in reverse lexicographic order (i.e reverse alphabetical order) by `cand_name`, along with their corresponding committee names. **Only select rows that have a matching `cand_id` in both tables.**

Your output should look similar to the following:

|cand_name|cmte_nm|
|----|----|
|ZUTLER, DANIEL PAUL MR|CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT|
|ZUMWALT, JAMES|ZUMWALT FOR CONGRESS|
|...|...|

Consider starting from the following query skeleton, which uses the `AS` keyword to rename the `cand` and `comm` tables to `c1` and `c2`, respectively.
Which join is most appropriate?

    SELECT ...
    FROM cand AS c1
        [INNER | {LEFT |RIGHT | FULL } {OUTER}] JOIN comm AS c2
        ON ...
    ...
    ...;


```python
# run this cell to make sure we're using the FEC dataset
%sql fec_engine 
```

```python
%%sql --save query_5a result_5a <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_5a = %sqlcmd snippets query_5a
save_results("result_5a", query_5a, result_5a)
result_5a
```

```python
grader.check("q5a")
```

<br/><br/>

---

### Question 5b

Suppose we modify the query from the previous part to include *all* candidates, **including those that don't have a committee.**


List the first 5 candidate names (`cand_name`) in reverse lexicographic order by `cand_name`, along with their corresponding committee names. If the candidate has no committee in the `comm` table, then `cmte_nm` should be NULL (or `None` in the `python` representation).

Your output should look similar to the following:

|cand_name|cmte_nm|
|----|----|
|ZUTLER, DANIEL PAUL MR|CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT|
|...|...|
|ZORNOW, TODD MR|None|

**Hint**: Start from the same query skeleton as the previous part. 
Which join is most appropriate?

```python
# run this cell to make sure we're using the FEC dataset
%sql fec_engine 
```

```python
%%sql --save query_5b result_5b <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_5b = %sqlcmd snippets query_5b
save_results("result_5b", query_5b, result_5b)
result_5b
```

```python
grader.check("q5b")
```

<br/><br/>

---

## Question 6: Subqueries and Grouping

If we return to our results from Question 4, we see that many of the contributions were to the same committee:

```python
# Your SQL query result from Question 4
# Reprinted for your convenience
result_4.DataFrame()['cmte_id'].value_counts()
```

<br><br>

For this question, create a new SQL query that returns the total amount that Donald Trump contributed to each committee.

Your table should have four columns: `cmte_id`, `total_amount` (total amount contributed to that committee), `num_donations` (total number of donations), and `cmte_nm` (name of the committee). Your table should be sorted in **decreasing order** of `total_amount`.

Your output should look similar to the following:

|cmte_id|total_amount|num_donations|cmte_nm|
|----|----|----|----|
|C00580100|18633157|131|DONALD J. TRUMP FOR PRESIDENT, INC.|
|C00055582|10000|1|NY REPUBLICAN FEDERAL CAMPAIGN COMMITTEE
|...|...|...|

**This is a hard question!** Don't be afraid to reference the lecture slides, or the overall SQL query skeleton at the top of this lab.

**Hint**:

* Note that committee names are not available in `indiv_sample_nyc`, so you will have to obtain information somehow from the `comm` table (perhaps a `JOIN` would be useful).
* Remember that you can compute summary statistics after grouping by using aggregates like `COUNT(*)`, `SUM()` as output fields.
* A **subquery** may be useful to break your question down into subparts. Consider the following query skeleton, which uses the `WITH` operator to store a subquery's results in a temporary table named `donations`.

        WITH donations AS (
            SELECT ...
            FROM ...
            ... JOIN ...
                ON ...
            WHERE ...
        )
        SELECT ...
        FROM donations
        GROUP BY ...
        ORDER BY ...;

```python
# run this cell to make sure we're using the FEC dataset
%sql fec_engine 
```

```python
%%sql --save query_6 result_6 <<
...
```

```python
# Run this cell before running the autograder; no further action is needed
query_6 = %sqlcmd snippets query_6
save_results("result_6", query_6, result_6)
result_6
```

```python
grader.check("q6")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Appa (left) and Momo (right) congratulate you for finishing Lab 10!

![](appa_momo.jpg) 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True, files=['results'])
```

Tests for q1 (points: None)
```
>>> query_q1, res_q1 = load_results('result_1') 
>>> res_q1
      titleType
0         short
1         movie
2      tvSeries
3       tvMovie
4  tvMiniSeries
5         video
6     videoGame
7     tvEpisode
8     tvSpecial
```
```
>>> query_q1, res_q1 = load_results('result_1') 
>>> res_q1.shape == (9, 1)
True
```
```
>>> query_q1, res_q1 = load_results('result_1') 
>>> "group by" not in query_q1.lower()
True
```
```
>>> query_q1, res_q1 = load_results('result_1') 
>>> sorted(list(res_q1.iloc[:,0])) == ['movie', 'short', 'tvEpisode', 'tvMiniSeries', 'tvMovie', 'tvSeries', 'tvSpecial', 'video', 'videoGame']
True
```


Tests for q2 (points: None)
```
>>> query_q2, res_q2 = load_results('result_2') 
>>> res_q2.shape == (12, 2)
True
```
```
>>> query_q2, res_q2 = load_results('result_2') 
>>> list(res_q2['total']) == [21665, 13830, 12175, 11028, 6995, 4123, 2747, 1558, 623, 410, 66, 6]
True
```
```
>>> query_q2, res_q2 = load_results('result_2') 
>>> list(res_q2['category']) == ['actor',
...  'writer',
...  'actress',
...  'producer',
...  'director',
...  'composer',
...  'cinematographer',
...  'editor',
...  'self',
...  'production_designer',
...  'archive_footage',
...  'archive_sound']
True
```


Tests for q3 (points: None)
```
>>> query_q3, res_q3 = load_results('result_3') 
>>> res_q3.shape == (10, 3)
True
```
```
>>> query_q3, res_q3 = load_results('result_3') 
>>> set(res_q3.columns) == set(['title', 'numVotes', 'averageRating'])
True
```
```
>>> query_q3, res_q3 = load_results('result_3') 
>>> np.sum(res_q3['numVotes'].astype(int)) == 19870486
True
```
```
>>> query_q3, res_q3 = load_results('result_3') 
>>> list(res_q3['title']) == ['The Shawshank Redemption',
...  'The Dark Knight',
...  'Inception',
...  'Fight Club',
...  'Pulp Fiction',
...  'Forrest Gump',
...  'Game of Thrones',
...  'The Matrix',
...  'The Lord of the Rings: The Fellowship of the Ring',
...  'The Lord of the Rings: The Return of the King']
True
```


Tests for q4 (points: None)
```
>>> query_q4, res_q4 = load_results('result_4') 
>>> res_q4.shape == (142, 3)
True
```
```
>>> query_q4, res_q4 = load_results('result_4') 
>>> sorted(res_q4['transaction_amt'].unique()) \
... == [224, 514, 1000, 1029, 1216, 1713, 2000, 2426, 2460, \
...     2529, 2531, 2536, 2544, 2574, 2587, 2600, 2620, 2679, \
...     2692, 2693, 2700, 2991, 3063, 3347, 3360, 3371, 3499, \
...     3548, 3706, 3954, 4049, 4127, 4164, 4591, 4931, 4958, \
...     5000, 5297, 5574, 6197, 6250, 6990, 7401, 7519, 7990, \
...     8735, 8896, 9000, 9013, 9106, 9651, 9752, 10000, 10418, \
...     11023, 15000, 23775, 2000000, 10000000]
True
```


Tests for q5a (points: None)
```
>>> query_q5a, res_q5a = load_results('result_5a') 
>>> res_q5a.shape == (5, 2)
True
```
```
>>> query_q5a, res_q5a = load_results('result_5a') 
>>> all(res_q5a == res_q5a.sort_values('cand_name', ascending=False))
True
```
```
>>> query_q5a, res_q5a = load_results('result_5a') 
>>> sorted(res_q5a['cmte_nm'].unique()) == [
...     'CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT',
...      'CONSTITUTIONAL COMMITTEE',
...      'JOE ZUCCOLO FOR CONGRESS',
...      'ZUKOWSKI FOR CONGRESS',
...      'ZUMWALT FOR CONGRESS']
True
```


Tests for q5b (points: None)
```
>>> query_q5b, res_q5b = load_results('result_5b') 
>>> res_q5b.shape == (5, 2)
True
```
```
>>> query_q5b, res_q5b = load_results('result_5b') 
>>> all(res_q5b == res_q5b.sort_values('cand_name', ascending=False))
True
```
```
>>> query_q5b, res_q5b = load_results('result_5b') 
>>> set(res_q5b['cand_name']) == set([
...     'ZUTLER, DANIEL PAUL MR', 
...     'ZUMWALT, JAMES',
...     'ZUKOWSKI, ANDREW GEORGE',
...     'ZUCCOLO, JOE',
...     'ZORNOW, TODD MR']) 
True
```
```
>>> query_q5b, res_q5b = load_results('result_5b') 
>>> sorted(map(str, res_q5b['cmte_nm'].unique())) == [
...     'CITIZENS TO ELECT DANIEL P ZUTLER FOR PRESIDENT',
...     'JOE ZUCCOLO FOR CONGRESS',
...     'None',
...     'ZUKOWSKI FOR CONGRESS',
...     'ZUMWALT FOR CONGRESS']
True
```
```
>>> query_q5b, res_q5b = load_results('result_5b') 
>>> res_q5b['cmte_nm'][4] == None
True
```


Tests for q6 (points: None)
```
>>> query_q6, res_q6 = load_results('result_6') 
>>> res_q6.shape == (10, 4)
True
```
```
>>> query_q6, res_q6 = load_results('result_6') 
>>> sorted(res_q6['total_amount']) == [1000, 2000, 2600, 5000, 5000, 5200, 5400, 9000, 10000, 18633157]
True
```
```
>>> query_q6, res_q6 = load_results('result_6') 
>>> sorted(res_q6['num_donations']) == [1, 1, 1, 1, 1, 1, 1, 2, 2, 131]
True
```
```
>>> query_q6, res_q6 = load_results('result_6') 
>>> sorted(res_q6['cmte_nm']) == [
...     'DONALD J. TRUMP FOR PRESIDENT, INC.',
...     'DONOVAN FOR CONGRESS',
...     'FRIENDS OF DAVE BRAT INC.',
...     'GRASSLEY COMMITTEE INC',
...     'HELLER FOR SENATE',
...     'NEW HAMPSHIRE REPUBLICAN STATE COMMITTEE',
...     'NY REPUBLICAN FEDERAL CAMPAIGN COMMITTEE',
...     'REPUBLICAN PARTY OF IOWA',
...     'SOUTH CAROLINA REPUBLICAN PARTY',
...     'TEXANS FOR SENATOR JOHN CORNYN INC']
True
```
